Before you turn this assignment in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All). Lastly, hit **Validate**.

If you worked locally, and then uploaded your work to the hub, make sure to follow these steps:
- open your uploaded notebook **on the hub**
- hit the validate button right above this cell, from inside the notebook

These  steps should solve any issue related to submitting the notebook on the hub.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Erica Chen"
COLLABORATORS = ""
```

---

# Lab 10: Feature Engineering & Cross-Validation
** This assignment is due 04/04/2018 at 11:59pm (graded on accuracy) **

In this lab, you will practice using scikit-learn to do feature engineering and cross-validation to produce a model with low error on held-out data.

```python
# Run this cell to set up your notebook
import seaborn as sns
import csv
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
%matplotlib inline
sns.set()
sns.set_context("talk")

from IPython.display import display, Latex, Markdown
```

### Introduction

For this lab, we will use a toy dataset to predict the house prices in Boston with data provided by the `sklearn.datasets` package.

Run the following cell to load the data. This will return a dictionary object which includes keys for:
    - `data` : the covariates (X)
    - `target` : the response vector (Y)
    - `feature_names`: the column names
    - `DESCR` : a full description of the data

```python
from sklearn.datasets import load_boston

boston_data = load_boston()
print(boston_data.keys())
```

Output:
```
dict_keys(['data', 'target', 'feature_names', 'DESCR'])
```

```python
print(boston_data['DESCR'])
```

Output:
```
Boston House Prices dataset
===========================

Notes
------
Data Set Characteristics:  

    :Number of Instances: 506 

    :Number of Attributes: 13 numeric/categorical predictive
    
    :Median Value (attribute 14) is usually the target

    :Attribute Information (in order):
        - CRIM     per capita crime rate by town
        - ZN       proportion of residential land zoned for lots over 25,000 sq.ft.
        - INDUS    proportion of non-retail business acres per town
        - CHAS     Charles River dummy variable (= 1 if tract bounds river; 0 otherwise)
        - NOX      nitric oxides concentration (parts per 10 million)
        - RM       average number of rooms per dwelling
        - AGE      proportion of owner-occupied units built prior to 1940
        - DIS      weighted distances to five Boston employment centres
        - RAD      index of accessibility to radial highways
        - TAX      full-value property-tax rate per $10,000
        - PTRATIO  pupil-
```

A look at the `DESCR` attribute tells us the data contains these features:

    1. CRIM      per capita crime rate by town
    2. ZN        proportion of residential land zoned for lots over 
                 25,000 sq.ft.
    3. INDUS     proportion of non-retail business acres per town
    4. CHAS      Charles River dummy variable (= 1 if tract bounds 
                 river; 0 otherwise)
    5. NOX       nitric oxides concentration (parts per 10 million)
    6. RM        average number of rooms per dwelling
    7. AGE       proportion of owner-occupied units built prior to 1940
    8. DIS       weighted distances to five Boston employment centres
    9. RAD       index of accessibility to radial highways
    10. TAX      full-value property-tax rate per 10,000 USD
    11. PTRATIO  pupil-teacher ratio by town
    12. B        1000(Bk - 0.63)^2 where Bk is the proportion of black 
                 residents by town
    13. LSTAT    % lower status of the population
    
Let's now convert this data into a pandas DataFrame. 

```python
boston = pd.DataFrame(boston_data['data'], columns=boston_data['feature_names'])
boston.head()
```

Output:
```
CRIM    ZN  INDUS  CHAS    NOX     RM   AGE     DIS  RAD    TAX  \
0  0.00632  18.0   2.31   0.0  0.538  6.575  65.2  4.0900  1.0  296.0   
1  0.02731   0.0   7.07   0.0  0.469  6.421  78.9  4.9671  2.0  242.0   
2  0.02729   0.0   7.07   0.0  0.469  7.185  61.1  4.9671  2.0  242.0   
3  0.03237   0.0   2.18   0.0  0.458  6.998  45.8  6.0622  3.0  222.0   
4  0.06905   0.0   2.18   0.0  0.458  7.147  54.2  6.0622  3.0  222.0   

   PTRATIO       B  LSTAT  
0     15.3  396.90   4.98  
1     17.8  396.90   9.14  
2     17.8  392.83   4.03  
3     18.7  394.63   2.94  
4     18.7  396.90   5.33
```

### Question 1

Let's model this housing price data! Before we can do this, however, we need to split the data into training and test sets. The latter, held-out points will be used to choose the best performing model. Remember that the response vector (housing prices) lives in the `target` attribute.

Use the [`train_test_split`](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.train_test_split.html) function to split out 10% of the data for test. Call the resulting splits `X_train`, `X_test`, `Y_train`, `Y_test`.

<!-- grade_id: q1_code, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import train_test_split
np.random.seed(47)

X = boston
Y = pd.Series(boston_data['target'])

X_train, X_test, Y_train, Y_test = train_test_split(test_size = 0.9, )
# YOUR CODE HERE
# raise NotImplementedError()
print(X_train.shape)
print(X_test.shape)
print(Y_train.shape)
print(Y_test.shape)
```

Output:
```
<ipython-input-7-2782a17a03fb> in <module>()
      5 Y = pd.Series(boston_data['target'])
      6 
----> 7 X_train, X_test, Y_train, Y_test = train_test_split(test_size = 0.9)
      8 # YOUR CODE HERE
      9 # raise NotImplementedError()

/Users/erica/anaconda/lib/python3.6/site-packages/sklearn/model_selection/_split.py in train_test_split(*arrays, **options)
   1675     n_arrays = len(arrays)
   1676     if n_arrays == 0:
-> 1677         raise ValueError("At least one array required as input")
   1678     test_size = options.pop('test_size', None)
   1679     train_size = options.pop('train_size', None)

ValueError: At least one array required as input
```

<!-- grade_id: q1_test, points: 1, grade: true, solution: false -->
```python
assert X_train.shape == (455, 13)
assert X_test.shape == (51, 13)
assert Y_train.shape == (455, )
assert Y_test.shape == (51, )
```

### Question 2

As a warmup, fit a linear model to describe the relationship between the housing price and all available covariates. We've imported `sklearn.linear_model` as lm, so you can use that instead of typing out the whole module name. Running the cell should create a scatter plot for our predictions vs the true prices.

<!-- grade_id: q2_code, points: 0, grade: false, solution: true -->
```python
import sklearn.linear_model as lm

linear_model = lm.LinearRegression()


# Fit your classifier
#linear_model.fit(...)

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: q2b_code, points: 0, grade: false, solution: true -->
```python

# Predict housing prices on the test set
Y_pred = ... # linear_model.predict(...)


# YOUR CODE HERE
raise NotImplementedError()

# Plot predicted vs true prices
plt.scatter(Y_test, Y_pred, alpha=0.5)
plt.xlabel("Prices")
plt.ylabel("Predicted prices")
plt.title("Prices vs Predicted prices")
```

### Question 3

As we find from the scatter plot, our model is not perfect. If it were perfect, we would see the identity line (i.e. a line of slope 1). Compute the root mean squared error (RMSE) of the predicted responses: 

$$
\textbf{RMSE} = \sqrt{\frac{1}{n}\sum_{i=1}^n \left( y_i - \hat{y}_i \right)^2 }
$$

Fill out the function below and compute the RMSE for our predictions on both the training data `X_train` and the test set `X_test`.  Note your implementation should not contain the word **"for"** (...that would be very slow).

<!-- grade_id: q3_code, points: 0, grade: false, solution: true -->
```python
def rmse(actual_y, predicted_y):
    """
    Args:
        predicted_y: an array of the prediction from the model
        actual_y: an array of the groudtruth label
        
    Returns:
        The root mean square error between the prediction and the groudtruth
    """
# YOUR CODE HERE
raise NotImplementedError()
    return ...
```

<!-- grade_id: q3_code2, points: 0, grade: false, solution: true -->
```python
train_error = ...
test_error = ...

# YOUR CODE HERE
raise NotImplementedError()

print("Training RMSE:", train_error)
print("Test RMSE:", test_error)
```

<!-- grade_id: q3_test, points: 1, grade: true, solution: false -->
```python
assert np.allclose((train_error, test_error), (4.56291225689, 5.88492861688))
```

## Cross Validation

**Warning**: don't use the test set to perform the feature selection! It may lead to over-fitting. We want to avoid using the test set too frequently. When selecting features or choosing hyper-parameters, we can split the training set further into train and validation sets. Then we can use the validation error to help select hyper-parameters.

Try $k$-fold cross-validation to select the best subset of features for our model. Recall the approach looks something like:

<img src="cv.png" width=500px>

### Question 4

Scikit-learn has built-in support for cross validation.  However, to better understand how cross validation works complete the following function which cross validates a given model.

1. Use the [`KFold.split`](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.KFold.html) function to get 4 splits on the training data. Note that `split` returns the indices of the data for that split.
2. For each split, select out the rows and columns based on the split indices and features.
3. Compute the RMSE on the validation split.
4. Return the average error across all cross validation splits.


<!-- grade_id: q4_code, points: 0, grade: false, solution: true -->
```python
from sklearn.model_selection import KFold

def compute_CV_error(model, X_train, Y_train):
    '''
    Split the training data into 4 subsets.
    For each subset, 
        fit a model holding out that subset
        compute the MSE on that subset (the validation set)
    You should be fitting 4 models total.
    Return the average MSE of these 4 folds.

    Args:
        model: an sklearn model with fit and predict functions 
        X_train (data_frame): Training data
        Y_train (data_frame): Label 

    Return:
        the average validation MSE for the 4 splits.
    '''
    kf = KFold(n_splits=4)
    validation_errors = []
    
    for train_idx, valid_idx in kf.split(X_train):
        # split the data
        split_X_train, split_X_valid = ..., ...
        split_Y_train, split_Y_valid = ..., ...

        # Fit the model on the training split
        ...
        
        # Compute the RMSE on the validation split
        error = ...

        # YOUR CODE HERE
        raise NotImplementedError()
        
        validation_errors.append(error)
        
    return np.mean(validation_errors)
```

<!-- grade_id: q5_test, points: 1, grade: true, solution: false -->
```python
assert np.isclose(
    compute_CV_error(lm.LinearRegression(), X_train[['TAX', 'INDUS', 'CRIM']],Y_train),
    7.5974094557701459)
```

### Question 5

Use the above cross validation function to determine which of the following feature sets to use:

<!-- grade_id: q5_code, points: 0, grade: false, solution: true -->
```python
feature_sets = [
    ['TAX', 'INDUS', 'CRIM'], 
    ['RM', 'LSTAT', 'PTRATIO'], 
    ['RM', 'B', 'NOX'], 
    ['TAX', 'LSTAT', 'DIS']
]

errors = []
for feat in feature_sets:
    print("Trying features:", feat)
    model = lm.LinearRegression()
    error = ... # compute the cross validation error
    print("\tRMSE:", error)
    errors.append(error)

best_err_idx = ...
best_err = ...
best_feature_set = ...

# YOUR CODE HERE
raise NotImplementedError()

for i in range(4):
    print('{}, error: {}'.format(feature_sets[i], errors[i]))

best_feature_set, best_err
```

<!-- grade_id: q5_test2, points: 1, grade: true, solution: false -->
```python
assert best_feature_set == ['RM', 'LSTAT', 'PTRATIO']
assert np.isclose(best_err, 5.221575997721903)
```

### Question 6
Finally, fit a linear classifier using your best feature set and predict housing prices for your original test set. Compute the final MSE.

<!-- grade_id: q6_code, points: 0, grade: false, solution: true -->
```python
# Fit your classifier
...

# Predict points from our test set and calculate the mse
train_rmse = ... 
test_rmse = ...


# YOUR CODE HERE
raise NotImplementedError()

print("Train RMSE", train_rmse)
print("KFold Validation RMSE", best_err)
print("Test RMSE", test_rmse)
```

Notice that the test error is higher than the validation error which is higher than the training error.  Why is this the case?

<!-- grade_id: q6_test, points: 1, grade: true, solution: false -->
```python
assert np.abs(test_rmse - 5.846401452163672) < 1e-3
```

Nice! You've used $k$-fold cross-validation to fit a linear regression model to the housing data.

In the future, you'd probably want to use something like [`cross_val_predict`](http://scikit-learn.org/stable/modules/generated/sklearn.model_selection.cross_val_predict.html) to automatically perform cross-validation, but it's instructive to do it yourself at least once.

## Submission

Congrats! You are finished with this assignment. Please don't forget to submit by 11:59 pm!